---
title: Hello IPL !!
description: Lets get started with Python

---
###  Python is a high-level programming language which means we need not to write the code in 0s and 1s. 

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: fd596ed3ea
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
