---
title       : Hello Cricket Lovers!!
description : Lets Get Started with Python.
---
## Hello World!!

One of the challenges of getting started with any programming language is to install the software and configure the path based on your operating system. Fortunately in this course we need to install anything and we can start coding in the browser. Thanks to DataCamp. 

Python is a high-level programming language which means we can write the programs in the english like language and Python interpreter will convert english like instructions into 0s and 1s. 

In Python we use `print` function for displaying data on the screen. In the right side window you can see 

```{python}
print("Hello Cricket Lovers")

```
It prints `Hello Cricket Lovers` to the screen. 

```yaml
type: NormalExercise
key: 42d1268a2b
lang: python
xp: 100
skills: 2
```


`@instructions`
use `print` fucntion to display `Hello World!!' 

`@hint`
you just need to replace Hello Cricket Lovers with Hello World!!

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
print("Hello Cricket Lovers")

```

`@solution`
```{python}
print("Hello World!!")
```

`@sct`
```{python}

```
