---
title       : Hello Cricket Lovers!!
description : Lets Get Started with Python.
---
## Hello World!!


```yaml
type: NormalExercise
key: 42d1268a2b
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
print("Hello Cricket Lovers")
for i in range(10):
    print(i)

```

`@solution`
```{python}

```

`@sct`
```{python}

```
